```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *
#import d8error
```

# Lab 3: Objects, Lists, Built-In Functions, Libraries

Welcome to Lab 3 of Introduction to Computational Thinking with Data! In this lab, we will be focusing on: 

<ul>
    <li> Debuggings code</li>
    <li> Continuing Objects and lists</li>
    <li> Getting familiar with Built-In Functions and Libraries</li>
    <li> Numpy Import Introduction</li>
</ul>

---

## Part 1: Comments and Debugging

<b>Debugging</b> in Python or any programming language involves identifying and fixing errors in your code, which can include syntax, semantic, and runtime issues.

<b> Comments </b> are hints that we add to our code to make it easier to understand. Python comments start with a <b>#</b>. Comments are completely ignored and not executed by code editors or interpreters.


### Debugging Examples

<div class="alert alert-info">
    <b>Fix the following Code bugs. Add a comment above the corrected code identifying the solution.</b>
</div>

```python
print "Hello"
```

Output:
```
Cell In[2], line 1
    print "Hello"
    ^
SyntaxError: Missing parentheses in call to 'print'. Did you mean print(...)?
```

```python
# type answer here
```

```python
print "Hello )
```

Output:
```
Cell In[3], line 1
    print "Hello )
          ^
SyntaxError: unterminated string literal (detected at line 1)
```

```python
# type answer here
```

```python
#expected outcome is 5
print("2" + "3")
```

Output:
```
23
```

```python
# type answer here
```

```python
#expected outcome is 2 + 3 = 5
print("2 + 3, 2 + 3)
```

Output:
```
Cell In[6], line 2
    print("2 + 3, 2 + 3)
          ^
SyntaxError: unterminated string literal (detected at line 2)
```

```python
# type answer here
```

```python
#expected output is 1
-x = 1
print (x)
```

Output:
```
Cell In[7], line 2
    -x = 1
    ^
SyntaxError: cannot assign to expression here. Maybe you meant '==' instead of '='?
```

```python
# type answer here
```

```python
#expected output 100
a = 10
print("a x 10", A)
```

```python
# type answer here
```

```python
a = [1,2,3
print a
```

Output:
```
Cell In[8], line 1
    a = [1,2,3
        ^
SyntaxError: '[' was never closed
```

```python
# type answer here
```

```python
sample = [25,98,632,,]
print(sampe)
```

Output:
```
Cell In[10], line 1
    sample = [25,98,632,,]
                        ^
SyntaxError: invalid syntax
```

```python
# type answer here
```

```python
sample = [-987, ”34 ” , 100, 900, 200, 50,987,1500]
print(sum( sample)
```

Output:
```
Cell In[12], line 1
    sample = [-987, ”34 ” , 100, 900, 200, 50,987,1500]
                    ^
SyntaxError: invalid character '”' (U+201D)
```

```python
# type answer here
```

---

## Python Library and Numpy Introduction

Before we get started here, make sure Numpy is installed on your machine or instance of Jupyterhub as this library is necessary for the following section.

<b>External libraries</b> in Python are Classes are collections of pre-written code that can used to solve programming tasks.

<b> import </b> in Python allows the programmer to access code defined in one module from another module.  

In the previous lab, we used the built-in function sum() to add values in a Python list. In the exercise below, we will use the Numpy .sum() to the values in the list.

#### Sum a List
Example:
```python
import numpy as np
list1 = [2,4,6,8]
x = np.sum(list1)
```

<div class="alert alert-info">
    <b>Write the code below: [ Find the sum using the Numpy .sum() ]</b>
</div>

```python
sample = [25, 34, 100, 900, 200,50]
print(sample)
```

```python
...
mean = (25 + 34 + 100 + 900 + 200 + 50)/ 6 # SOLUTION
```

```python
grader.check("ex1")
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for ex1 (points: 2)
```
>>> type(mean) == float
True
```
```
>>> mean - (25 + 34 + 100 + 900 + 200 + 50) / 6 < 0.0001
True
```
